rule_files:
  - "../*.alerts.yml"

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'up{job="prometheus", instance="localhost:9090"}'
        values: "0 0 0 0 0 0 1 0 0 0 0 1 0 0 0"
      - series: 'up{job="node_exporter", instance="localhost:9100"}'
        values: "1+0x6 0 0 0 0 0 0 1 0" # 1 1 1 1 1 1 1 0 0 0 0 0 0 1 0
      - series: 'go_goroutines{job="prometheus", instance="localhost:9090"}'
        values: "10+10x2 30+20x5" # 10 20 30 30 50 70 90 110 130
      - series: 'go_goroutines{job="node_exporter", instance="localhost:9100"}'
        values: "10+10x7 10+30x4" # 10 20 30 40 50 60 70 80 10 40 70 100 130

    alert_rule_test:
      - eval_time: 5m
        alertname: InstanceDown
        exp_alerts:
          - exp_labels:
              severity: page
              instance: localhost:9090
              job: prometheus
            exp_annotations:
              summary: "Instance localhost:9090 down"
              description: "localhost:9090 of job prometheus has been down for more than 5 minutes."
      - eval_time: 12m
        alertname: InstanceDown
        exp_alerts:
          - exp_labels:
              severity: page
              instance: localhost:9100
              job: node_exporter
            exp_annotations:
              summary: "Instance localhost:9100 down"
              description: "localhost:9100 of job node_exporter has been down for more than 5 minutes."
      - eval_time: 7m
        alertname: GoroutineHigh
        exp_alerts:
          - exp_labels:
              severity: critical
              instance: localhost:9090
              job: prometheus
            exp_annotations:
              summary: "Goroutine high"
              description: "Goroutine high (current value: 110)"
          - exp_labels:
              severity: critical
              instance: localhost:9100
              job: node_exporter
            exp_annotations:
              summary: "Goroutine high"
              description: "Goroutine high (current value: 80)"
    promql_expr_test:
      - expr: go_goroutines > 5
        eval_time: 4m
        exp_samples:
          - labels: 'go_goroutines{job="prometheus",instance="localhost:9090"}'
            value: 50
          - labels: 'go_goroutines{job="node_exporter",instance="localhost:9100"}'
            value: 50
      - expr: sum_over_time(up[10m])
        eval_time: 15m
        exp_samples:
          - labels: '{job="node_exporter",instance="localhost:9100"}'
            value: 3
          - labels: '{job="prometheus",instance="localhost:9090"}'
            value: 2
      - expr: sum_over_time(up[10m]) / count_over_time(up[10m])
        eval_time: 15m
        exp_samples:
          - labels: '{job="node_exporter",instance="localhost:9100"}'
            value: 0.3
          - labels: '{job="prometheus", instance="localhost:9090"}'
            value: 0.2
